# SLU08 - Functions Intermediate: Exercise notebook

```python
# Make sure to run this cell or the grading won't work!

import json
import hashlib
import inspect
import numpy as np
from random import randint
from utils import check_call
```

## Exercise 1 (ungraded): Scope Basics

Before diving into code, let’s make sure you understand some key terms. In your own words, describe what the following concepts mean:

1. **Scope**  
2. **Namespace**  
3. **Local Scope**  
4. **Global Scope**

Take a moment to reflect and write your answers below.

---

💡 *Hint:* If you’re unsure, revisit the Learning Notebook section on **Scopes and Namespaces**. These are important ideas that will save you a lot of debugging time later on!

📌 This is an ungraded question, but it's a very important one. Feel free to discuss your definitions with the other students or instructors in Slack!

1. scope 

## Exercise 2: Scopes and Namespace

Consider the code snippet below. Your job is to **identify the scope** of each variable.

```python
def welcome_user():
    a_var = 10
    return a_var

another_var = 100
```

**Your Task:**
Assign one of the following strings to each answer:
- `local` → The variable is defined inside a function and only exists there.
- `global` → The variable is defined at the top level and accessible anywhere in the script.

*Think about where each variable is created and where it's accessible.*

Fill in the variables:

<!-- grade_id: cell-954725562970a8a3, points: 0, grade: false, solution: true -->
```python
a_var_scope = 
another_var_scope = 
```

<!-- grade_id: cell-5beb29cb41412c43, points: 2, grade: true, solution: false -->
```python
assert hashlib.sha256(a_var_scope.encode()).hexdigest() == '25bf8e1a2393f1108d37029b3df5593236c755742ec93465bbafa9b290bddcf6', 'Incorrect answer for a_var_scope, try again!'
assert hashlib.sha256(another_var_scope.encode()).hexdigest() == '8001c27439650c5c5a6b4ed94163b5ddeb4476362c71380e613fa20dfffcef50', 'Incorrect answer for another_var_scope, try again!'
print('Very good, your answer is correct!')
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In[14], line 1
----> 1 assert hashlib.sha256(a_var_scope.encode()).hexdigest() == '25bf8e1a2393f1108d37029b3df5593236c755742ec93465bbafa9b290bddcf6', 'Incorrect answer for a_var_scope, try again!'
      2 assert hashlib.sha256(another_var_scope.encode()).hexdigest() == '8001c27439650c5c5a6b4ed94163b5ddeb4476362c71380e613fa20dfffcef50', 'Incorrect answer for another_var_scope, try again!'
      3 print('Very good, your answer is correct!')

AttributeError: 'int' object has no attribute 'encode'
```

### Exercise 3: What's going on with scopes?
Consider the code below and try to determine what will be printed by the four print statements. **Do this without running the code** to test your understanding of variable scopes.

```python
def parallel_universe():
    a = 10
    print(a)  # first print

    def hello_universe():
        print(a)  # second print

    def good_morning_sunshine(a):
        print(a)  # third print
        a = 50
    
    hello_universe()
    good_morning_sunshine(a)

a = 200
parallel_universe()
print(a)  # fourth print
```

<!-- grade_id: cell-f108f9b2b1624025, points: 0, grade: false, solution: true -->
```python
# print_1 =
# print_2 =
# print_3 =
# print_4 =

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-6f0bede927d6d49f, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(str(print_1).encode()).hexdigest() == '4a44dc15364204a80fe80e9039455cc1608281820fe2b24f1e5233ade6af1dd5', 'print_1 is not correct.'
assert hashlib.sha256(str(print_2*2).encode()).hexdigest() == 'f5ca38f748a1d6eaf726b8a42fb575c3c71f1864a8143301782de13da2d9202b', 'print_2 is not correct.'
assert hashlib.sha256(str(print_3*3).encode()).hexdigest() == '624b60c58c9d8bfb6ff1886c2fd605d2adeb6ea4da576068201b6c6958ce93f4', 'print_3 is not correct.'
assert hashlib.sha256(str(print_4*4).encode()).hexdigest() == '1a1cf797fabe7f95836fabeca626907c77b3e6c9aff7c2290b396a238c69362e', 'print_4 is not correct.'
print('Very good, your answer is correct!')
```

### Exercise 4:  Fancy Return Statement
Write a function called calculate_stats that takes a list of numbers and returns both the sum and the average of the list, as a tuple.

If you're feeling fancy, try writing the return statement on a single line 😉

<!-- grade_id: cell-bdc3f99859e53dcd, points: 0, grade: false, solution: true -->
```python
# Modify this function:
def calculate_stats(numbers):
# YOUR CODE HERE
raise NotImplementedError()

stats_sum, stats_avg = calculate_stats([1, 2, 3])
```

<!-- grade_id: cell-16341d4e8ac60c98, points: 2, grade: true, solution: false -->
```python
assert isinstance(stats_sum, int) and isinstance(stats_avg, float), "The returned values should be an int and a float."

result_hash = hashlib.sha256(json.dumps((stats_sum, stats_avg)).encode()).hexdigest()
expected_hash = "7665a2b8868d014556eab0bf2dce92924595c18e770ff58450960f7802b8061b"

assert result_hash == expected_hash, "Oops! The returned values are not quite right."

print("Great job! Your function returns the correct sum and average.")

```

## Exercise 5: Default Arguments
The function `calculate_total_price` below is used to compute the final price of a product, taking into account a discount and tax rate. Your job is to:
- Add default values to the discount and `tax_rate` parameters:
    - The `discount` should default to `0.`
    - The `tax_rate` should default to `0.23` (the default VAT in Portugal at the time of writing).
- Test the function with both default and custom values.

**TASK**: Modify this function by adding default values

```python
def calculate_total_price():
    final_price = price - discount
    final_price *= (1 + tax_rate)
    return round(final_price, 2)
```

<!-- grade_id: cell-eb4e9cfc837086c7, points: 0, grade: false, solution: true -->
```python
# Write the complete def calculate_total_price():
# YOUR CODE HERE
raise NotImplementedError()

# Run this cell to test your function
price_only = calculate_total_price(100)
custom = calculate_total_price(100, discount=10, tax_rate=0.2)

print("Price with defaults:", price_only)
print("Price with custom values:", custom)
```

<!-- grade_id: cell-6938960a9a0f275f, points: 3, grade: true, solution: false -->
```python
hash_price_only = hashlib.sha256(str(price_only).encode()).hexdigest()
hash_custom = hashlib.sha256(str(custom).encode()).hexdigest()

assert hash_price_only == "0589abecb964e139e06fd64a22365653e56ca8fb6d83226d4e395cac3de94b49", \
    "Check the formula and default values for price_only."

assert hash_custom == "c0faa5f07166b1dfd7e5250aaedaaff87d3b9c989a010d57c28d80fe43b5e521", \
    "Custom discount and tax_rate calculation isn't correct."

print("Well done! Your default arguments and calculations are correct.")

```

### Exercise 6: Test your Scope knowledge
Create a function with the following specifications:
- Create a function called `use_global_scope` that takes two arguments
- The first argument is called `multiply_by` and should be an integer
- The second argument is called `a_number` and is a number
- The function should return the value of `a_number` multiplied by `multiply_by`
- If `multiply_by` is not an integer, the function returns `None`

<!-- grade_id: cell-677ed0e5bc6a02ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eb4e97d47bbb5aa8, points: 4, grade: true, solution: false -->
```python
# check function source
sig = inspect.signature(use_global_scope)
source = inspect.getsource(use_global_scope)
source_lines = source.split('\n')
params = sig.parameters

# check function name
assert ' '.join(source_lines[0].split('(')[0].split())=='def use_global_scope', 'The function name is not correct.' 

# check function parameters
assert len(params) == 2, 'The function should take two parameters.'
assert list(params.keys())[0] == 'multiply_by', 'The name of the first parameter is not correct.'
assert list(params.keys())[1] == 'a_number', 'The name of the second parameter is not correct.'

# test function
i1=randint(0,100)
i2=randint(0,100)
assert use_global_scope(i1,i2) == i1*i2, 'The function is not multiplying correctly.'
i1=randint(0,100)
i2=randint(0,100)/4
assert use_global_scope(i1,i2) == i1*i2, 'The function is not multiplying correctly.'
assert use_global_scope(1.45,6739) == None, 'If multiply_by is not an integer, the function should return None.'

print('Very good, your answer is correct!')
```

### Exercise 6.2 
Now use the function you've defined in 6.1:
- Create a global variable called `a_number` and assign it the value of `2`.
- Call the `use_global_scope` function, passing one positional and one keyword argument.
- Use the global variable as the value of the `a_number` argument.
- Use 14 as the value of the `multiply_by` argument.

<!-- grade_id: cell-d77f348198eb2cd6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ea37f9eaa3577b0, points: 3, grade: true, solution: false -->
```python
# check the a_number variable
'a_number' in globals().keys(), 'The variable a_number does not exist.'
assert a_number == 2, 'The value of a_number should be 2.'

# check the function call
keys_list=list(globals().keys())
l_keys=len(keys_list)
output=' '.join([str(globals()[keys_list[i]]) \
                 for i in range(l_keys-10,l_keys)])
output_w_spaces=''.join(output.split())
assert check_call(output_w_spaces), 'Are you calling the function as '\
    'specified?'

print('Very good, your answer is correct!')
```

## Exercise 7: Wrap-Up Exercise: Smart Calculator

🚀 Your First Smart Calculator!

Functions don't just automate repetitive tasks — they allow us to **build tools** that make coding easier, faster, and more fun! 😎

In this challenge, you're going to create your very own **calculator function**. But not just any calculator — a smart one that explains what it's doing.

### Your task:
Write a function called `calculator` that:
- Accepts **two required arguments**: `num1` and `num2` (both numbers).
- Accepts an **optional** third argument: `operator`, which defaults to `"+"`.
- Returns a **tuple** with:
  1. the calculated result
  2. a **description string** like `"5 + 2 = 7"`

You should support the following operators: `"+"`, `"-"`, `"*"`, `"/"`.

If an unsupported operator is passed in, return:
- `"Invalid operator"` as the result
- A description string like `"5 ? 2 = Invalid operator"` (already given)

You got this! 💪

<!-- grade_id: cell-e22f24afa7857f62, points: 0, grade: false, solution: true -->
```python
# Create a calculator function that:
# - Accepts two numbers
# - Has an optional operator argument (default: "+")
# - Returns both the result and a description string like "5 + 2 = 7"

def calculator(num1, num2, operator="+"):
    # YOUR CODE HERE
    raise NotImplementedError()
    description = f"{num1} {operator} {num2} = {result}"
    
    return result, description
    
calc_result = calculator(10, 2, "/")
print(calc_result)
```

<!-- grade_id: cell-7e07021dfb20db78, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(str(calc_result).encode()).hexdigest() == "723e087e68c944b5111c9894f54e31a0c27a37e202917b12116979e3642cc855", "Check your division logic and the format of the returned string."

print("Well done! Your calculator is computing and describing the operation correctly.")

```

# Submit your work!

To submit your work, [follow the steps here, in the step "Grading the Exercise Notebook"!](https://github.com/LDSSA/ds-prep-course-2025/blob/main/docs/weekly-workflow.md#13-work-on-the-learning-units)